<?xml version="1.0" encoding="UTF-8"?>
<quiz>
    <question type="stack">
        <name>
            <text>Randomized Integral Evaluation and Area Under Curve</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[<p>For the function f(x) = {@a@}x^3 - {@b@}x^2 + {@c@}x - {@d@}, please complete the following tasks:</p>
      <p><b>Part 1:</b> Evaluate the indefinite integral of f(x).</p>
      <p><b>Part 2:</b> Using your result from Part 1, find the definite integral of f(x) from x=0 to x=2.</p>
      <p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[<p>This question checks your understanding of basic antiderivatives and the Power Rule for Integration, the Fundamental Theorem of Calculus, the geometric interpretation of integrals, and your competency in mathematical communication and problem-solving.</p>]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.3333333</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2021080100</text>
        </stackversion>
        <questionvariables>
            <text>a:rand([1,2,3,4,5,6]);b:rand([1,2,3,4,5,6]);c:rand([1,2,3,4,5,6]);d:rand([1,2,3,4,5,6])</text>
        </questionvariables>
        <specificfeedback format="html">
            <text>[[feedback:prt1]]</text>
        </specificfeedback>
        <questionnote>
            <text></text>
        </questionnote>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>1</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
        </prtincorrect>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>asin</inversetrig>
        <matrixparens>(</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>int(a*x^3-b*x^2+c*x-d,x)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords></forbidwords>
            <allowwords></allowwords>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options></options>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackvariables></feedbackvariables>
            <node>
                <name>0</name>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>int(a*x^3-b*x^2+c*x-d,x)</tans>
                <testoptions></testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0.0000000</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Your integration is correct!]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0.1000000</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Your integration is not correct. Remember to use the Power Rule for Integration.]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
</quiz>
